<problem url_name="problem_bd3b45d8bb48" markdown="null">
  <script type="text/python" system_path="python_lib">
from mitxgraders import *
import sys
syspath = "\n".join(sys.path)
grader = StringGrader(
    answers='cat'
)
          </script>
  <customresponse cfn="grader" expect="cat">
    <textline/>
    <pre>$syspath</pre>
  </customresponse>
</problem>
